# Problem 3

Gwyn hates words with vowels. Help Gwyn design a function that takes one parameter: `message`. The function should remove all the vowels from `message` then return the modified vowel-less `message`.

Note: You need to use a for/while loop for this problem.

### Examples:
```python
remove_vowels("the quick brown fox jumps over the lazy dog") == "th qck brwn fx jmps vr th lzy dg"
remove_vowels("one night and one more time, thanks for the memories") == "n nght nd n mr tm, thnks fr th mmrs"
```

### Assumptions:
- `y` isn't a vowel for the purposes of this problem.
- All letters will be in lowercase.
- Spaces and punctuation marks should remain intact.


<!-- grade_id: cell-ee24ab0c5a330f58, points: 0, grade: false, solution: true -->
```python
def remove_vowels(message):
    vowels = 'aeiou'
    out = ''
    for i in message:
        if i in vowels:
            out += ''
        else:
            out += i
    return out
```

<!-- grade_id: cell-71a9b48bfa6316df, points: 0, grade: true, solution: false -->
```python
from utils.utils import print_and_assert

print_and_assert( remove_vowels, "th qck brwn fx jmps vr th lzy dg", "the quick brown fox jumps over the lazy dog" )
print_and_assert( remove_vowels, "n nght nd n mr tm, thnks fr th mmrs", "one night and one more time, thanks for the memories" )
print_and_assert( remove_vowels, "", "" )
print_and_assert( remove_vowels, "ppl", "apple" )

```

Output:
```
Test Case with parameters ('the quick brown fox jumps over the lazy dog',)
Expected output: th qck brwn fx jmps vr th lzy dg
Your output: th qck brwn fx jmps vr th lzy dg
Success!

Test Case with parameters ('one night and one more time, thanks for the memories',)
Expected output: n nght nd n mr tm, thnks fr th mmrs
Your output: n nght nd n mr tm, thnks fr th mmrs
Success!

Test Case with parameters ('',)
Expected output: 
Your output: 
Success!

Test Case with parameters ('apple',)
Expected output: ppl
Your output: ppl
Success!
```

```python

```