# Problem 4

You are planning your annual field trip for your family this weekend. It will be a road trip visiting different places across Luzon. One important thing you need to consider is whether you have enough fuel for your trip. Design a function that takes two parameters: `itinerary` and `liters_fuel_left`. The function should calculate and return whether it's possible to complete the trip given `liters_fuel_left`.

- Your car uses 3L per 10km.
- The possible return values are:
    - `"You do not have enough fuel for your trip. You need __ more liters."` where __ is the number of liters you still need.
    - `"You have enough fuel for your trip."`
- `itinerary` is a list of positive integers that represent the distance (in km) you need to travel between places. The distances will be multiples of 10.
- `liters_fuel_left` is the amount of fuel left.

Note: You are required to use a for/while loop for this problem.

For example:

Let's say you are going to Town A which takes 70km, then travel to Town B which takes 100km, then to Town C which takes 30km, then back home which takes 20km. You would need 66L of fuel to complete the whole trip.

### Examples:
```python
check_fuel([70, 100, 30, 20], 70) == "You have enough fuel for your trip."
check_fuel([70, 100, 30, 20], 60) == "You do not have enough fuel for your trip. You need 6.0 more liters."
check_fuel([70, 100, 30, 20], 66) == "You have enough fuel for your trip."

check_fuel([50, 60], 35) == "You have enough fuel for your trip."
check_fuel([50, 20, 60], 40) == "You have enough fuel for your trip."

check_fuel([50, 60], 10) == "You do not have enough fuel for your trip. You need 23.0 more liters."
check_fuel([50, 20, 60], 38) == "You do not have enough fuel for your trip. You need 1.0 more liters."
```

### Assumptions:
- All the distances will be multiples of 10.
- You won't be stopping on ANY gas stations during the trip.
- `liters_fuel_left` is a positive integer.


<!-- grade_id: cell-5083630e5a7cbe80, points: 0, grade: false, solution: true -->
```python
def check_fuel(itinerary, liters_fuel_left):
    total_km = 0
    for i in itinerary:
        total_km += i
        fuel_needed = (total_km / 10) * 3
    if fuel_needed > liters_fuel_left:
        fuel_needed = fuel_needed - liters_fuel_left
        return "You do not have enough fuel for your trip. You need " + str(fuel_needed) + " more liters."
    else:
        return "You have enough fuel for your trip."
```

<!-- grade_id: cell-383ae4d2b188bdcb, points: 10, grade: true, solution: false -->
```python
from utils.utils import print_and_assert

print_and_assert( check_fuel, "You have enough fuel for your trip.", [70, 100, 30, 20], 70 )
print_and_assert( check_fuel, "You do not have enough fuel for your trip. You need 6.0 more liters.", [70, 100, 30, 20], 60 )
print_and_assert( check_fuel, "You have enough fuel for your trip.", [70, 100, 30, 20], 66 )

print_and_assert( check_fuel, "You have enough fuel for your trip.", [50, 60], 35 )
print_and_assert( check_fuel, "You have enough fuel for your trip.", [50, 20, 60], 40 )

print_and_assert( check_fuel, "You do not have enough fuel for your trip. You need 23.0 more liters.", [50, 60], 10 )
print_and_assert( check_fuel, "You do not have enough fuel for your trip. You need 1.0 more liters.", [50, 20, 60], 38 )

```

Output:
```
Test Case with parameters ([70, 100, 30, 20], 70)
Expected output: You have enough fuel for your trip.
Your output: You have enough fuel for your trip.
Success!

Test Case with parameters ([70, 100, 30, 20], 60)
Expected output: You do not have enough fuel for your trip. You need 6.0 more liters.
Your output: You do not have enough fuel for your trip. You need 6.0 more liters.
Success!

Test Case with parameters ([70, 100, 30, 20], 66)
Expected output: You have enough fuel for your trip.
Your output: You have enough fuel for your trip.
Success!

Test Case with parameters ([50, 60], 35)
Expected output: You have enough fuel for your trip.
Your output: You have enough fuel for your trip.
Success!

Test Case with parameters ([50, 20, 60], 40)
Expected output: You have enough fuel for your trip.
Your output: You have enough fuel for your trip.
Success!

Test Case with parameters ([50, 60], 10)
Expected output: You do not have enough fuel for your trip. You need 23.0 more liters.
Your outpu
```

```python

```