# Problem 1

Your dear uncle is not getting enough sleep lately. He has asked you for some help with his insomnia. You decided to design an LED display that will display a message that would let him count some sheep. Design a function that takes one parameter: `n`. The function should **return** a string that counts sheep up to `n`. For example, if `n` is 4, then the function should return `"1 sheep...2 sheep...3 sheep...4 sheep..."`.

Note: You need to use a for/while loop for this problem.

### Examples:
```python
count_sheep(4) == "1 sheep...2 sheep...3 sheep...4 sheep..."
count_sheep(6) == "1 sheep...2 sheep...3 sheep...4 sheep...5 sheep...6 sheep..."
```

### Assumptions:
- `n` is a positive integer


<!-- grade_id: cell-c267f0b6c88922a6, points: 0, grade: false, solution: true -->
```python
def count_sheep(n):
    out = ''
    for i in range(1, n + 1):
        result = str(i) + " sheep..."
        out += result
    return out
```

<!-- grade_id: cell-f549d262406b272d, points: 12, grade: true, solution: false -->
```python
from utils.utils import print_and_assert

print_and_assert( count_sheep, "1 sheep...2 sheep...3 sheep...4 sheep...", 4 )
print_and_assert( count_sheep, "1 sheep...2 sheep...3 sheep...4 sheep...5 sheep...6 sheep...", 6 )

```

Output:
```
Test Case with parameters (4,)
Expected output: 1 sheep...2 sheep...3 sheep...4 sheep...
Your output: 1 sheep...2 sheep...3 sheep...4 sheep...
Success!

Test Case with parameters (6,)
Expected output: 1 sheep...2 sheep...3 sheep...4 sheep...5 sheep...6 sheep...
Your output: 1 sheep...2 sheep...3 sheep...4 sheep...5 sheep...6 sheep...
Success!
```

```python

```