# Problem 2

Design a function that takes four parameters: `x1`, `y1`, `x2`, `y2`. These parameters represent two points: `(x1, y1)` and `(x2, y2)`. The function should return a string whose value depends on whether the line the two points form is `increasing`, `decreasing`, or `horizontal`.

Remember that:
- An `increasing` line has a positive slope. It looks like this:

<img src="increasing line.png" width="250"/>

- A `decreasing` line has a negative slope. It looks like this:

<img src="decreasing line.png" width="250"/>

- A `horizontal` line has a slope that is 0.

- The formula for the slope `m` is as follows:

<img src="slope.png" width="200" />


### Examples:
```python
line_direction(-4, -2, 3, 3) == "increasing"
line_direction(-1, 5, -4, 0) == "increasing"
line_direction(-4, 3, -2, 0) == "decreasing"
line_direction(2, 8, 6, 1) == "decreasing"
line_direction(2, 2, -7, 2) == "horizontal"
```

### Assumptions:
- The line will never be a vertical line (because it's slope would be undefined).
- The two points will never coincide.


<!-- grade_id: cell-9aa5037c377b6819, points: 0, grade: false, solution: true -->
```python
def line_direction(x1, y1, x2, y2):
    p1 = y2 - y1
    p2 = x2 - x1
    m = p1 / p2
    if m > 0:
        return "increasing"
    elif m == 0:
        return "horizontal"
    return "decreasing"

```

<!-- grade_id: cell-f554b99971457b2e, points: 0, grade: true, solution: false -->
```python
from utils.utils import print_and_assert

print_and_assert( line_direction, "increasing", -4, -2, 3, 3 )
print_and_assert( line_direction, "increasing", -1, 5, -4, 0 )
print_and_assert( line_direction, "decreasing", -4, 3, -2, 0 )
print_and_assert( line_direction, "decreasing", 2, 8, 6, 1 )
print_and_assert( line_direction, "horizontal", 2, 2, -7, 2 )

```

Output:
```
Test Case with parameters (-4, -2, 3, 3)
Expected output: increasing
Your output: increasing
Success!

Test Case with parameters (-1, 5, -4, 0)
Expected output: increasing
Your output: increasing
Success!

Test Case with parameters (-4, 3, -2, 0)
Expected output: decreasing
Your output: decreasing
Success!

Test Case with parameters (2, 8, 6, 1)
Expected output: decreasing
Your output: decreasing
Success!

Test Case with parameters (2, 2, -7, 2)
Expected output: horizontal
Your output: horizontal
Success!
```

```python

```